---
rule_files:
  - ../test-application-ingress.yaml

evaluation_interval: 1h

tests:
  - interval: 1h
    input_series:
      - series: 'nginx_ingress_controller_requests{exported_namespace="test-application-dev", ingress="test-application-v1-2", path="/", status="200"}'
        values: '0 1 3 2+4x10 5 6'
      - series: 'nginx_ingress_controller_requests{exported_namespace="test-application-dev", ingress="test-application-v1-2", path="/", status="201"}'
        values: '0 2 5 2+4x10 5 6'

    name: 2xxResponses alert fired as not enough requests in given time window

    promql_expr_test:
      - expr: sum by (ingress, exported_namespace) (increase(nginx_ingress_controller_requests{exported_namespace="test-application-dev", ingress="test-application-v1-2", path="/", status=~"2.*"}[1h])) < 10
        eval_time: 2h
        exp_samples:
          - labels: '{exported_namespace="test-application-dev", ingress="test-application-v1-2"}'
            value: 5 # this is (5-2)+(3-1)
      - expr: sum by (ingress, exported_namespace) (increase(nginx_ingress_controller_requests{exported_namespace="test-application-dev", ingress="test-application-v1-2", path="/", status=~"2.*"}[1h])) < 10
        eval_time: 8h
        exp_samples:
          - labels: '{exported_namespace="test-application-dev", ingress="test-application-v1-2"}'
            value: 8 # since incrementing every hour by 4 for each ingress
      - expr: sum by (ingress, exported_namespace) (increase(nginx_ingress_controller_requests{exported_namespace="test-application-dev", ingress="test-application-v1-2", path="/", status=~"2.*"}[1h])) < 10
        eval_time: 14h
        exp_samples: [] # no samples since we have 5+5 requests at this time

    alert_rule_test:
      # assert that the 2xx responses alert is firing since less than default of 100 for last 60 minutes
      - eval_time: 12h
        alertname: 2xxResponses
        exp_alerts:
          - exp_labels:
              alertname: 2xxResponses
              exported_namespace: test-application-dev
              ingress: test-application-v1-2
              severity: alert-severity-tag
              application: test-application
              businessUnit: hmpps
              environment: none
              productId: none
            exp_annotations:
              message: "Ingress test-application-dev/test-application-v1-2 is not serving enough 2xx responses."
              runbook_url: "https://runbook.com/ingress-2xx-responses"
              dashboard_url: "https://grafana.live.cloud-platform.service.justice.gov.uk/d/f1e13059dfd23fdcaf479f4fa833f92610c2dfa5/kubernetes-ingress-traffic?1&var-namespace=test-application-dev&var-ingress=test-application(-v1-2)&from=now-2h&to=now"
